---
title: Next Steps
description: |-
  - quick summary of what's been learned
  - the Unix philosophy: lots of little tools combined in ad hoc ways
  - Noble's Rules for laying out projects (and why)
  - pointer to the second course (loops, finding files, etc.)

--- type:NormalExercise lang:python xp:100 skills:2 key:cd80efde24
## Tidying Up

Using whatever commands you want, reorganize this project to follow the structure laid out in Noble's Rules. Run the command `bin/check` to see what's out of place.

*** =instructions

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}

```

*** =solution
```{python}

```

*** =sct
```{python}

```
