---
title       : Chapter 1
description : This is a chapter description.

--- type:MultipleChoiceExercise lang:python xp:50 skills:2 key:d17c2c5b68
## This is a multiple choice

Here's some assignment text. What's your name?

*** =instructions
- Greg
- Nick
- Kara
- Hugo

*** =hint
A hint here.

*** =pre_exercise_code
```{python}

```

*** =sct
```{python}

```
