---
title: Core Commands
description: >-
  - `whoami`

  - `pwd` and a comment on how cryptic commands are

  - interpreting output of `pwd`: folders containing folders

  - `ls` (depends on a certain file setup)

  - `cat`

  - Editing text files (use `nano` and explain how a shell can have multiple
  personalities)

  - Text vs. binary files

  - `cp`

  - `mv`

  - `rm`

  - `mkdir` and `rmdir`

--- type:MultipleChoiceExercise lang:python xp:50 skills:2 key:ba5ac63f56
## Absolute and Relative Paths

Starting from `/Users/jasmine/data/`, which of the following commands could Jasmine use to navigate to  `/Users/jasmine`?

1. `cd .`
2. `cd /`
3. `cd /home/jasmine`
4. `cd ../..`
5. `cd ~`
6. `cd home`
7. `cd ~/data/..`
8. `cd`
9. `cd ..`

*** =instructions

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sct
```{python}

```
