---
title: Why the Command Line?
description: >-
  Three-minute video explaining that the command line is an interface to the
  operating system, just like the GUIs that learners are probably familiar with,
  but because it's textual, it supports history and automation.

--- type:MultipleChoiceExercise lang:python xp:50 skills:2 key:ae26c736be
## How are they connected?

What is the relationship between the graphical file explorer that most people use and the shell?
- The file explorer lets you view and edit files, while the shell lets you run programs.
- The file explorer is built on top of the shell.
- The shell is part of the operating system, while the file explorer is separate.
- They are both interfaces for issuing commands to the operating system.

*** =instructions

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sct
```{python}

```

--- type:MultipleChoiceExercise lang:python xp:50 skills:2 key:6eee4e6c00
## What is a shell?

What kind of thing is a shell?
- a program
- an operating system
- a text-only mode of the file explorer

*** =instructions

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sct
```{python}

```
